---
## Q2 (2 points)
จงเขียนโปรแกรมที่รับชื่อและนามสกุลจากผู้ใช้ แล้วนำมาสร้างเป็นสตริงใหม่ 2 สตริง โดย

- สตริงหนึ่งเป็นชื่อต่อด้วยนามสกุลและมีช่องว่าง (blank) คั่นกลาง

- อีกสตริงหนึ่งเป็นนามสกุลต่อด้วยชื่อและมี , ต่อด้วยช่องว่างคั่นกลาง

แล้วแสดงผลลัพธ์ดังตัวอย่างข้างล่างนี้

### ตัวอย่างการทำงาน: 

Firstname Lastname: Somsak Saksri

Hello! Somsak Saksri

Saksri, Somsak 

### ตัวอย่างการทำงาน: 

Firstname Lastname: Jackson Wang

Hello! Jackson Wang

Wang, Jackson 

<!-- grade_id: name, points: 0, grade: false, solution: true -->
```python
def run():
    # YOUR CODE HERE
    
```

```python
'''ตรวจสอบผลลัพธ์ของโปรแกรมด้วยการรันที่ cell นี้'''
run()
```

<!-- grade_id: correct_name, points: 2, grade: true, solution: false -->
```python
'''ตรวจสอบคำสั่ง print และผลลัพธ์'''
from unittest.mock import patch
from io import StringIO

def test(inputs,expected_output): 
    with patch('builtins.input', side_effect=inputs), patch('sys.stdout', new_callable=StringIO) as mock_stdout:
        run()
        assert mock_stdout.getvalue() == expected_output,mock_stdout.getvalue()
        
test(['Somsak Saksri'],'Hello! Somsak Saksri\nSaksri, Somsak\n')
test(['Jackson Wang'],'Hello! Jackson Wang\nWang, Jackson\n')
```