---
## Q3 (2 points)
จงเขียนโปรแกรมที่รับชื่อและอายุจากผู้ใช้ แล้วคำนวณปีเกิดของผู้ใช้และแสดงผลลัพธ์ดังตัวอย่างข้างล่างนี้

### ตัวอย่างการรัน:

Name, age: Somsak 19 

Somsak's year of birth is 2546.

### ตัวอย่างการรัน:

Name, age: Bella 2 

Bella's year of birth is 2563.

<!-- grade_id: year_of_birth, points: 0, grade: false, solution: true -->
```python
def run():
    # YOUR CODE HERE
    
```

```python
'''ตรวจสอบผลลัพธ์ของโปรแกรมด้วยการรันที่ cell นี้'''
run()
```

<!-- grade_id: correct_year_of_birth, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบการคำนวณปีเกิด'''
from unittest.mock import patch
from io import StringIO
from datetime import date 

def test(inputs,expected_output,exact_output=True): 
    with patch('builtins.input', side_effect=inputs), patch('sys.stdout', new_callable=StringIO) as mock_stdout:
        run()     
        if exact_output:
            assert mock_stdout.getvalue() == expected_output,mock_stdout.getvalue()
        else:
            assert expected_output in mock_stdout.getvalue(),mock_stdout.getvalue()
        
test(['Somsak 19'],'2546',exact_output=False)
test(['Bella 2'],'2563',exact_output=False)
```

<!-- grade_id: exact, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบการคำนวณปีเกิดและการแสดงผลลัพธ์'''
test(['Somsak 19'],"Somsak's year of birth is 2546.\n",exact_output=True)
test(['Bella 2'],"Bella's year of birth is 2563.\n",exact_output=True)
test(['Somsak 69'],"Somsak's year of birth is 2496.\n",exact_output=True)
test(['Kim 65'],"Kim's year of birth is 2500.\n",exact_output=True)
```