---
## Q4 (4 points)
จงเขียนโปรแกรมที่รับชื่อและอายุจากผู้ใช้ แล้วคำนวณปีเกิดของผู้ใช้เพื่อสร้าง email address ของผู้ใช้โดย
เอาชื่อมาต่อด้วยเลข 2 หลักหลังของปีเกิดและต่อด้วย @chula.ac.th แสดงผลลัพธ์ดังตัวอย่างข้างล่างนี้

### ตัวอย่างการทำงาน: 

Name: Somsak 

Age : 19 

Somsak46@chula.ac.th

### ตัวอย่างการทำงาน: 
        
Name: Jaidee

Age : 64 

Jaidee01@chula.ac.th

### ตัวอย่างการทำงาน: 
        
Name: Yaya

Age : 70

Yaya95@chula.ac.th

<!-- grade_id: email_address, points: 0, grade: false, solution: true -->
```python
def run():
    # YOUR CODE HERE
```

```python
'''ตรวจสอบผลลัพธ์ของโปรแกรมด้วยการรันที่ cell นี้'''
run()
```

<!-- grade_id: exact, points: 3, grade: true, solution: false -->
```python
from unittest.mock import patch
from io import StringIO
from datetime import date 

def test(inputs,expected_output): 
    with patch('builtins.input', side_effect=inputs), patch('sys.stdout', new_callable=StringIO) as mock_stdout:
        run()
        assert mock_stdout.getvalue() == expected_output,mock_stdout.getvalue()
        
'''ตรวจสอบเลขสองตัวท้ายและการแสดงผลลัพธ์'''
'''กรณีเลขสองตัวท้ายของปีเกิดอยู่ในช่วง 10-99'''
test(['Somsak','19'],'Somsak46@chula.ac.th\n')
test(['Jaidee','64'],'Jaidee01@chula.ac.th\n')  
test(['Yaya','70'],'Yaya95@chula.ac.th\n')  

```

<!-- grade_id: twodigit, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบเลขสองตัวท้ายและการแสดงผลลัพธ์'''
'''กรณีเลขสองตัวท้ายของปีเกิดอยู่ในช่วง 00-09'''
test(['Somsak','65'],'Somsak00@chula.ac.th\n')
test(['Yaya','56'],'Yaya09@chula.ac.th\n')  
```