## Q3 (3 points)
จงเขียนโปรแกรมที่รับความยาวด้าน 3 ด้านของสามเหลี่ยมจากผู้ใช้ ซึ่งค่าที่รับมาจะมากกว่าศูนย์เสมอ แล้วตรวจสอบว่าเป็นด้านของสามเหลี่ยม
หรือไม่ และแสดงผลลัพธ์ดังตัวอย่างข้างล่างนี้

###  ตัวอย่างการทำงาน: 

Length of 3 sides: **3, 7.5, 4** <br>
Triangle: False

###  ตัวอย่างการทำงาน: 

Length of 3 sides: **3.4, 5, 4** <br>
Triangle: True

<!-- grade_id: triangle, points: 0, grade: false, solution: true -->
```python
def run():
    # YOUR CODE HERE
    a, b, c = [float(x) for x in (input("Length of 3 sides: ").split(","))]
    if 0 not in [a, b, c]:
        print("Triangle:", a + b > c and a + c > b and b + c > a)
```

```python
'''ตรวจสอบผลลัพธ์ของโปรแกรมด้วยการรันที่ cell นี้'''
run()
```

Output:
```
Length of 3 sides: 3,4,5
Triangle: True
```

<!-- grade_id: partial, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณ'''
from unittest.mock import patch
from io import StringIO
import math

def test(inputs,expected_output,exact_output): 
    with patch('builtins.input', side_effect=inputs), patch('sys.stdout', new_callable=StringIO) as mock_stdout:
        run()
        if exact_output:
            assert mock_stdout.getvalue() == expected_output,mock_stdout.getvalue()
        else:
            assert expected_output in mock_stdout.getvalue(),mock_stdout.getvalue()
test(['3,7.5,4'],'False',exact_output=False)
test(['3.4,5,4'],'True',exact_output=False)            
            
```

<!-- grade_id: exact, points: 2, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
test(['3,7.5,4'],'Triangle: False\n',exact_output=True)
test(['3.4,5,4'],'Triangle: True\n',exact_output=True)

```

```python

```

Output:
```
False
```

```python

```