## Q1 (2 points)
จงเขียนโปรแกรมที่คำนวณหาผลเฉลยของสมการกำลังสอง (quadratic equation) โดยให้รับสัมประสิทธิ์
$𝑎, 𝑏, 𝑐$ ในสมการ $𝑎𝑥^2 + 𝑏𝑥 + 𝑐 = 0$ จากผู้ใช้แล้วคำนวณหาผลเฉลยจากสูตร $𝑥 = \frac{−𝑏± \sqrt{b^2-4ac}} {2a} $ แล้ว
แสดงผลลัพธ์ดังตัวอย่างข้างล่างนี้ ($𝑎$ ต้องไม่เป็น $0$ และ $𝑏^2 > 4𝑎𝑐$)

หมายเหตุ: การหารากที่สองทำได้โดยใช้การยกกำลัง ½ เช่น 2**0.5 หรือ ใช้ฟังก์ชัน sqrt ในโมดูล math ซึ่ง
ต้องใช้คำสั่ง import math เพื่อระบุว่าโปรแกรมนี้ใช้ฟังก์ชันในโมดูล math และ เรียกใช้ฟังก์ชัน sqrt ดัง
ตัวอย่างนี้

import math

z = math.sqrt(2) # value in z is 1.4142135…

### ตัวอย่างการทำงาน: 

Enter coefficients a, b, c : **1,4,3** <br>
x = -1.0 , -3.0

###  ตัวอย่างการทำงาน: 

Enter coefficients a, b, c : **2.5,-8.7,2.1** <br>
x = 3.219053751558745 , 0.26094624844125425

<!-- grade_id: quardatic, points: 0, grade: false, solution: true -->
```python
def run():
    # YOUR CODE HERE
    a, b, c = [float(x) for x in (input("Enter coefficients a, b, c : ").split(","))]
    if a == 0 or b ** 2 < 4 * a * c:
        return
    top_right = (b ** 2) - (4 * a * c)
    x1 = (-b + (top_right ** 0.5)) / (2 * a)
    x2 = (-b - (top_right ** 0.5)) / (2 * a)
    print("x = {} , {}".format(x1, x2))
```

```python
'''ตรวจสอบผลลัพธ์ของโปรแกรมด้วยการรันที่ cell นี้'''
run()
```

Output:
```
Enter coefficients a, b, c : 1,4,3
x = -1.0 , -3.0
```

<!-- grade_id: partial, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณ'''
from unittest.mock import patch
from io import StringIO
import math

def test(inputs,expected_output,exact_output): 
    with patch('builtins.input', side_effect=inputs), patch('sys.stdout', new_callable=StringIO) as mock_stdout:
        run()
        if exact_output:
            x, eq, r1, comma, r2 = expected_output.split(' ')
            r1 = round(float(r1),4)
            r2 = round(float(r2),4)
            x_ans, eq_ans, ans1, ans_comma, ans2 = mock_stdout.getvalue().split(' ')
            ans1 = round(float(ans1),4)
            ans2 = round(float(ans2),4)
            assert r1==ans1 and comma==ans_comma and r2==ans2 , mock_stdout.getvalue()
        else:
            r1, comma, r2 = expected_output.split(' ')
            r1 = math.floor(float(r1))
            r2 = math.floor(float(r2))
            ans = mock_stdout.getvalue().split(' ')
            ans1 = math.floor(float(ans[-3]))
            ans2 = math.floor(float(ans[-1]))
            assert r1==ans1 and r2==ans2, mock_stdout.getvalue()
       
test(['1,4,3'],'-1.0 , -3.0',exact_output=False)
test(['2.5,-8.7,2.1'],'3.219053751558745 , 0.26094624844125425',exact_output=False)


```

<!-- grade_id: exact, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
test(['1,4,3'],'x = -1.0 , -3.0\n',exact_output=True)
test(['2.5,-8.7,2.1'],'x = 3.219053751558745 , 0.26094624844125425\n',exact_output=True)

```

```python

```