## Q2 (2 points)
จงเขียนโปรแกรมที่หาว่าสามารถหาผลเฉลยของสมการกำลังสองจากสูตร $𝑥 = −𝑏±\sqrt{b^2-4ac}$ ได้หรือไม่ 

โดยตรวจสอบว่า $𝑏^2 − 4𝑎𝑐 ≥ 0$ และ $𝑎 ≠ 0$ แล้วแสดงผลลัพธ์ดังตัวอย่างข้างล่างนี้

###  ตัวอย่างการทำงาน: 

Enter coefficients a, b, c : **1, 4, 3** <br>
Can use quadratic formula: True

###  ตัวอย่างการทำงาน: 

Enter coefficients a, b, c : **2.3, -1.5, 1.5** <br>
Can use quadratic formula: False

<!-- grade_id: quardatic, points: 0, grade: false, solution: true -->
```python
def run():
    # YOUR CODE HERE
    a, b, c = [float(x) for x in (input("Enter coefficients a, b, c : ").split(","))]
    if a == 0 or b ** 2 < 4 * a * c:
        print("Can use quadratic formula:",False)
    else:
        print("Can use quadratic formula:",True)
```

```python
'''ตรวจสอบผลลัพธ์ของโปรแกรมด้วยการรันที่ cell นี้'''
run()
```

Output:
```
Enter coefficients a, b, c : 1,4,3
Can use quadratic formula: True
```

<!-- grade_id: partial, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณ'''
from unittest.mock import patch
from io import StringIO
import math

def test(inputs,expected_output,exact_output): 
    with patch('builtins.input', side_effect=inputs), patch('sys.stdout', new_callable=StringIO) as mock_stdout:
        run()
        if exact_output:
            assert mock_stdout.getvalue() == expected_output,mock_stdout.getvalue()
        else:
            assert expected_output in mock_stdout.getvalue(),mock_stdout.getvalue()
            
test(['1,4,3'],str(True),exact_output=False)
test(['2.5,-1.5,1.5'],str(False),exact_output=False)   

```

<!-- grade_id: exact, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
test(['1,4,3'],'Can use quadratic formula: True\n',exact_output=True)
test(['2.5,-1.5,1.5'],'Can use quadratic formula: False\n',exact_output=True)

```

```python

```